Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `???`, `YOUR CODE HERE`, "???", "YOUR ANSWER HERE", as well as your name and collaborators below:

```scala
val NAME = ""
val COLLABORATORS = ""
```

---

# Exercise: Expressions and Data Types

<!-- 3 Expressions -->

<!-- 4 Binding and Scope -->

The purpose of this assignment is to warm-up with Scala.

### Learning Goals

The primary learning goals of this assignment are to build intuition for
the following:

-   thinking in terms of types, values, and expressions; and
-   imperative iteration.

### Instructions

This assignment asks you to write Scala code. There are restrictions
associated with how you can solve these problems. Please pay careful
heed to those. If you are unsure, ask the course staff.

Note that `???` indicates that there is a missing function or code
fragment that needs to be filled in. In Scala, it is also an expression
that throws a `NotImplementedError` exception. Make sure that you remove
the `???` and replace it with the answer.

Use the test cases provided to test your implementations. You are also
encouraged to write your own test cases to help debug your work.
However, please delete any extra cells you may have created lest they
break an autograder.

## Type Checking

In the following, I have left off the return type of function `g`. The
body of `g` is well-typed if we can come up with a valid return type. In
this question, we will reason for ourselves that `g` is indeed
well-typed.

```scala
def g(x: Int) = /*e1*/({
  // env1
  val (a, b) = /*e2*/(
    // env2
    (1, (x, 3))
  )
  /*e3*/(
    // env3
    if (x == 0) (b, 1) else (b, a + 2)
  )
})
```

We have added parentheses around 3 key sub-expressions of the body of
`g` (e.g., `/*e1*/(`$\ldots$`)`) and noted that there are corresponding
environments (e.g., `// env1` for each sub-expression).

### Q1.1: What is the type environment `env1`? Briefly explain.

Use either format shown in the notes (e.g.,
$[ x_1 : \tau_1, \ldots, x_n : \tau_n ]$).

<!-- grade_id: q1_1, points: 2, grade: true, solution: true -->
???

### Q1.2: What is the type environment `env2`? Briefly explain.

<!-- grade_id: q1_2, points: 2, grade: true, solution: true -->
???

### Q1.3: Derive the type of expression `e2`.

Showing the type for each sub-expression of `e2`; stop when you reach
literals or variable uses.

<!-- grade_id: q1_3, points: 10, grade: true, solution: true -->
???

#### Notes

Use the format shown in the notes. Here’s such a derivation for the type
of the expression `(1 + 2) + (3 + 4)`.

``` scala
(1 + 2) + (3 + 4): Int
- if 1 + 2: Int
  - if 1: Int
  - if 2: Int
- if 3 + 4: Int
  - if 3: Int
  - if 4: Int
```

### Q1.4: What is the type environment `env3`? Briefly explain.

<!-- grade_id: q1_4, points: 2, grade: true, solution: true -->
???

### Q1.5: Derive the type of expression `e3`.

<!-- grade_id: q1_5, points: 2, grade: true, solution: true -->
???

### Q1.6: Confirm your derivations by adding type assertions to each sub-expression of `g` and adding the return type of `g`.

That is, replace sub-expressions $e$ of `g` with expressions $e$ `:`
$\tau$. You may need to add some parentheses—`(` $e$ `:` $\tau$ `)`—to
preserve the syntactic structure. Skip adding typing assertions for
literals and variable uses.

**Edit this cell:**

<!-- grade_id: q1_6, points: 9, grade: true, solution: true -->
```scala
def g(x: Int) = /*e1*/({
  // env1
  val (a, b) = /*e2*/(
    // env2
    (1, (x, 3))
  )
  /*e3*/(
    // env3
    if (x == 0) (b, 1) else (b, a + 2)
  )
})
```

*Hint*: There are 8 sub-expressions that are not literals nor variable
uses that need typing assertions (i.e., $e$ `:` $\tau$), plus 1 more
annotation for the return type of `g`.

## Unit Testing

When starting to program in the large, it is useful to use a testing
framework to manage tests and integrate with IDEs. One that is commonly
used in Scala is [ScalaTest](https://www.scalatest.org/).

While it is somewhat overkill for testing small exercises like the ones
to come, we practice here writing tests using ScalaTest.

To load the ScalaTest library, run the following cell:

```scala
// RUN this cell FIRST before testing
import $ivy.`org.scalatest::scalatest:3.2.19`, org.scalatest._, events._, flatspec._
def report(suite: Suite) = suite.execute(stats = true)
def assertPassed(suite: Suite) =
  suite.run(None, Args(new Reporter {
    def apply(e: Event) = e match {
      case e @ (_: TestFailed) => assert(false, s"${e.message} (${e.testName})")
      case _ => ()
    }
  }))
def test(suite: Suite) = {
  report(suite)
  assertPassed(suite)
}
```

### Q2.1: Unit Test `plus`

For this question, **edit the next two code cells** to fix the
implementation of `plus` and add the appropriate assertion for the third
test case `"add (3,4) == 7"`.

Our goal is unit test the following complicated function (that we’ve
gotten wrong!):

```scala
def plus(n: Int, m: Int): Int = 2
```

To use ScalaTest, we create “Spec” objects using ScalaTest methods like
`should` that define an embedded domain-specific language (DSL) for
defining tests:

```scala
val plusSuite = new AnyFlatSpec {
  // Define a *subject* to test (e.g., "plus").
  // After `should`, name a test (e.g, "add (1, 1) == 2").
  // After `in`, specify assertions (e.g., `assert(plus(1,1) == 2))
  "plus" should "add 1 + 1 == 2" in {
    // Specify assertions here.
    assert(plus(1,1) == 2)
  }

  it should "add 2 + 2 == 4" in {
    // It is convenient to distinguish the expected result from the code that
    // you're testing, which affects the error messages when the test fails.
    assertResult(2 + 2) {
      plus(2,2)
    }
  }

  it should "add 3 + 4 == 7"  in {
    ???
  }
}
report(plusSuite)
```

<!-- grade_id: q2_1, points: 3, grade: true, solution: false -->
```scala
assertPassed(plusSuite)
```

## Run-Time Library

Most languages come with a standard library with support for things like
data structures, mathematical operators, string processing, etc.
Standard library functions may be implemented in the object language
(perhaps for portability) or the meta language (perhaps for
implementation efficiency).

For this question, we will implement some library functions in Scala,
our meta language, that we can imagine will be part of the run-time for
our object language interpreter. In actuality, the main purpose of this
exercise is to warm-up with Scala programming.

### Q3.1: Write and test a function `abs`

**Edit this cell:**

<!-- grade_id: q3_1_code, points: 0, grade: false, solution: true -->
```scala
def abs(n: Double): Double = ???
```

that returns the absolute value of `n`. This a function that takes a
value of type `Double` and returns a value of type `Double`. This
function corresponds to the JavaScript library function `Math.abs`.

#### Notes

-   Do not use any Scala library functions.

#### Tests

**Edit this cell:**

```scala
val absSuite = new AnyFlatSpec {
  "abs" should "abs(2) == 2" in {
     assert(abs(2) == 2)
  }
  it should "abs(-2) == 2" in {
     assert(abs(-2) == 2)
  }
  it should "abs(0) == 0" in {
     assert(abs(0) === 0)
  }
  it should "???1" in {
    ???
  }
}
report(absSuite)
```

<!-- grade_id: q3_1, points: 4, grade: true, solution: false -->
```scala
assertPassed(absSuite)
```

### Q3.2: Write and test a function `xor`

**Edit this cell:**

<!-- grade_id: q3_2_code, points: 0, grade: false, solution: true -->
```scala
def xor(a: Boolean, b: Boolean): Boolean = ???
```

that returns the exclusive-or of `a` and `b`. The exclusive-or returns
`true` if and only if exactly one of `a` or `b` is `true`.

#### Notes

-   For practice, do not use the Boolean operators. Instead, only use
    the `if`- expression and the Boolean literals (i.e., `true` or
    `false`).

#### Tests

**Edit this cell:**

```scala
val xorSuite = new AnyFlatSpec {
  "xor" should "!xor(true, true)" in {
    assert(!xor(true, true))
  }
  it should "xor(true, false)" in {
    assert(xor(true, false))
  }
  it should "???1" in {
    ???
  }
  it should "???2" in {
    ???
  }
}
report(xorSuite)
```

<!-- grade_id: q3_2, points: 4, grade: true, solution: false -->
```scala
assertPassed(xorSuite)
```

## Imperative Iteration and Complexity

### Q4.1: Write a function `filterPairsByBound`

**Edit this cell:**

<!-- grade_id: q4_1_code, points: 0, grade: false, solution: true -->
```scala
def filterPairsByBound(list: List[(Int, Int)], k: Int): List[(Int, Int)] =
  ???
```

that given a list of pairs of integers, for example,

```scala
val input1_l = List( (1, 5), (2, 7), (15, 14), (18, 19), (14, 28), (0,0), (35, 24) )
```

output a list consisting of just those pairs $(n_1, n_2)$ in the
original list wherein $|n_1 - n_2 \leq k$ where $k$ is an integer given
as input. Ensure that the order of the elements in the output list is
the same as that in the input list.

For the list `input1`, the expected output, with `k == 1`, the expected
output is as follows:

```scala
val input1_k = 1
val output1 = List( (15,14), (18,19), (0,0) )
assert(filterPairsByBound(input1_l, input1_k) == output1)
```

With `k == 4`, the expected output is as follows:

```scala
val input2_k = 4
val output2 = List( (1, 5), (15, 14), (18, 19), (0,0) )
assert(filterPairsByBound(input1_l, input2_k) == output2)
```

#### Notes

-   Your function must be called `filterPairsByBound` with two
    arguments: (1) a list of pairs of integers, and (2) the $k$ value.
    It must return a list of pairs of integers.
-   You can use `for`-loops (or `foreach`) and the following operators
    for concatenating elements to a list:
    -   `:+` appends an element to the back of a list.
    -   `:::` concatenates two lists together.
    -   `::` prepends an element to the front of a list.
-   You can use the `List` API method `reverse`. You may also use the
    `Int` `abs` method to obtain the absolute value of an integer (or
    use your `abs` function above).
-   You should not use any other `List` API functions including
    `filter`, `map`, `foldLeft`, `foldRight`, etc. Plenty of time to
    learn them properly later on.
-   Do not try to convert your list to an array or vector so that you
    can mutate it.
-   If you are unable to solve the problem without violating the
    restrictions or unsure, ask us.
-   You will need to use `var` given the above restrictions.

#### Hints

-   In Scala, pairs of integers have the type `(Int, Int)`.
-   A list containing pairs of integers has the type `List[(Int, Int)]`.
-   Recall from the notes, here is how one iterates over the elements of
    a list in Scala:

```scala
val list = List(1, 2, 3)

for (elt <- list) {
  // do stuff with elt
  println(elt)
}
```

-   Append an element to the end of a list and update a `var`:

```scala
var resultList = List(1, 2, 3)
val elt = 42

resultList = resultList :+ elt
```

-   Or, append an element to the end of a list using list concatenation
    and update a `var`:

```scala
var resultList = List(1, 2, 3)
val elt = 42

resultList = resultList ::: List(elt)
```

-   Prepend an element and update a `var`:

```scala
var resultList = List(1, 2, 3)
val elt = 42

resultList = elt :: resultList
```

-   *Warning*: The `:+`, `:::`, or other operations appending operations
    take linear $O(n)$ time where $n$ is the length of the (left) list.
    Thus, we will often try to avoid using these operations, but it is
    ok to use it for this particular part.

#### Tests

<!-- grade_id: q4_1, points: 5, grade: true, solution: false -->
```scala
test(new AnyFlatSpec {
  "filterPairsByBound" should "test1" in {
    val lst1 = List((1,1), (-1,1), (1,-1))
    assertResult( List((1,1)) ){ filterPairsByBound(lst1,1) }
   }

  it should "test2" in {
    val lst2 = List((1,2), (2,1), (1,0), (0,1), (1,1))
    assertResult( lst2 ){ filterPairsByBound(lst2,1) }
  }

  it should "test3" in {
    val lst3 = List((1,5), (5,1))
    assertResult( Nil ){ filterPairsByBound(lst3,1) }
  }

  it should "test4" in {
    val lst4 = List((2,5), (1,2), (-5,-4), (4,1))
    assertResult( List((1,2), (-5,-4)) ){ filterPairsByBound(lst4,1) }
  }

  it should "test5" in {
    val lst5 = List( (1, 5), (2, 7), (15, 14), (18, 19), (14, 28), (0,0), (35, 24) )
    assertResult( List((1, 5), (15, 14), (18, 19), (0,0)) ){ filterPairsByBound(lst5,4) }
  }

  it should "test6" in {
    val lst6 = List((1, 2), (2, 1), (1,0), (0,1), (1,1))
    assertResult( List((1,1)) ){ filterPairsByBound(lst6,0) }
  }
})
```

### Q4.2: Write a function `filterPairsByBoundLinearTime`

**Edit this cell:**

<!-- grade_id: q4_2_code, points: 0, grade: false, solution: true -->
```scala
def filterPairsByBoundLinearTime(list: List[(Int, Int)], k: Int): List[(Int, Int)] =
  ???
```

If you followed the hint and ignored the linear-time warning in the
previous part
<a href="#sec-filterPairsByBound" class="quarto-xref">Section 4.1</a>,
then you would have used the `:+` or `:::` operation to append an
element to the end of a list at each step.

``` scala
for (... <- list) {
  // iterate over a loop
  ...
  newList = newList :+ newElement // This takes O(length of newList).
}
```

Each `:+` or `:::` operation requires a full list traversal to find the
end of `newList` and then append to it. The overall algorithm thus
requires $O(n^2)$ time where is the length of the original `list` (also
the number of loop iterations).

To illustrate, cut-and-paste and then run this code in a new test cell.
Just remember to delete that cell before you submit. It will take a long
time to run.

``` scala
// Create a list of 1,000,000 pairs
val longTestList = (1 to 1000000).map(x => (x, x - 1)).toList
// Run the function you wrote
filterPairsByBound(longTestList, 1)
// This will take a long time to finish.
```

In this problem, we wish to implement a function
`filterPairsByBoundLinearTime` that solves the exact same problem as the
previous part
<a href="#sec-filterPairsByBound" class="quarto-xref">Section 4.1</a>
but takes time linear in the size of the input list.

To do so, we would like you to use the `::` (read “cons”) operator on a
list that prepends an element to the front of a list, instead of `:+` or
`:::` that appends to the back of a list.

You will want to use the `List` `reverse` API method:

```scala
val list = List(1, 2, 5, 6, 7, 8)
val r = list.reverse
```

The `r` has the reverse of `list`, and it works in linear time in the
length of `list`.

The restrictions remain the same as the previous part
<a href="#sec-filterPairsByBound" class="quarto-xref">Section 4.1</a>,
but we would like you to focus on ensuring that your solution runs in
linear time.

#### Tests

<!-- grade_id: q4_2, points: 7, grade: true, solution: false -->
```scala
test(new AnyFlatSpec {
  "filterPairsByBoundLinearTime" should "test1" in {
    val lst1 = List((1,1), (-1,1), (1,-1))
    assertResult( List((1,1)) ){ filterPairsByBoundLinearTime(lst1,1) }
   }

  it should "test2" in {
    val lst2 = List((1,2), (2,1), (1,0), (0,1), (1,1))
    assertResult( lst2 ){ filterPairsByBoundLinearTime(lst2,1) }
  }

  it should "test3" in {
    val lst3 = List((1,5), (5,1))
    assertResult( Nil ){ filterPairsByBoundLinearTime(lst3,1) }
  }

  it should "test4" in {
    val lst4 = List((2,5), (1,2), (-5,-4), (4,1))
    assertResult( List((1,2), (-5,-4)) ){ filterPairsByBoundLinearTime(lst4,1) }
  }

  it should "test5" in {
    val lst5 = List( (1, 5), (2, 7), (15, 14), (18, 19), (14, 28), (0,0), (35, 24) )
    assertResult( List((1, 5), (15, 14), (18, 19), (0,0)) ){ filterPairsByBoundLinearTime(lst5,4) }
  }

  it should "test6" in {
    val lst6 = List((1, 2), (2, 1), (1,0), (0,1), (1,1))
    assertResult( List((1,1)) ){ filterPairsByBoundLinearTime(lst6,0) }
  }
})
```

## Submission

#### Submission Instructions

If you are a University of Colorado Boulder student, we use Gradescope
for assignment submission. In summary,

-   [ ] Work on a copy of this Jupyter notebook.
-   [ ] Submit it to the corresponding Gradescope assignment entry for
    grading.

#### GitHub and Gradescope

We use GitHub Classroom for assignment distribution, which gives you a
private GitHub repository to work on your assignment. While using GitHub
is perhaps overkill for this assignment, it does give you the ability to
version and save incremental progress on GitHub (lest your laptop fails)
and makes it easier to get help from the course staff. It will also
become particularly useful when more files are involved, and it is never
too early to get used to the workflow professional software engineers
use with Git.

To use use GitHub and Gradescope,

-   [ ] Create a private GitHub repository by clicking on a GitHub
    Classroom link from the corresponding Canvas assignment entry.
-   [ ] Clone your private GitHub repository to your development
    environment (using the `<> Code` button on GitHub to get the
    repository URL).
-   [ ] Work on the copy of this Jupyter notebook from your cloned
    repository. Use Git to save versions on GitHub (e.g., `git add`,
    `git commit`, `git push` on the command line, via JupyterLab, or via
    VSCode).
-   [ ] Submit to the corresponding Gradescope assignment entry for
    grading by choosing GitHub as the submission method.

You need to have a GitHub identity and must have your full name in your
GitHub profile in case we need to associate you with your submissions.